Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Abinesh Senthil Kumar"
COLLABORATORS = ""
```

---

<!-- grade_id: cell-7736b26b66387e08, points: 0, grade: false, solution: true -->
```python
# Exercise 1
def sphere_area(radius):
    # Write a function that returns the surface area of a sphere given a radius. 
    # Use the input `radius` as the radius
    # round the surface area to two decimales 
    # Use the PI from the math module
    
    import math
    PI = math.pi
    surfacearea = 4 * PI * (radius**2)
    return round(surfacearea,2)
    
```

<!-- grade_id: cell-cd169ff880b9c9af, points: 2, grade: true, solution: false -->
```python
# Test cell #1 for exercise 1
assert sphere_area(15) == 2827.43
assert sphere_area(7.25) == 660.52
```

<!-- grade_id: cell-96d4c8eab6a5297c, points: 2, grade: true, solution: false -->
```python
# Test cell for exercise 1 (hidden)
```

<!-- grade_id: cell-fc31ef88e4698e4c, points: 0, grade: false, solution: true -->
```python
# Exercise 2
def sphere_volume(radius):
    # Write a function that returns the volume of a sphere given a radius. 
    # Use the input `radius` as the radius
    # round the surface area to two decimales 
    # Use the PI from the math module
    
    import math
    PI = math.pi
    
    volume = (4/3) * PI * radius**3
    return round(volume,2)
```

<!-- grade_id: cell-7afc760eadc1b572, points: 2, grade: true, solution: false -->
```python
# Test cell #1 for exercise 2
assert sphere_volume(15) == 14137.17
assert sphere_volume(7.5) == 1767.15
```

<!-- grade_id: cell-5ba5475eb0d9a795, points: 2, grade: true, solution: false -->
```python
# Test cell for exercise 2 (Hidden)
```

<!-- grade_id: cell-f7c6618b722b51d0, points: 0, grade: false, solution: true -->
```python
# Exercise 3
def total_wages_for_the_week(hours, wage):
    # Many companies pay time-and-a-half for any hours worked above 40 in a given week.
    # Write a function that uses the inputs to this function, hours worked (hours) and the hourly rate (wage),  and 
    # calculate the total wages for the week. 
    # YOUR CODE HERE
    if hours <= 40:
        wage1 = hours * wage
        return wage1
    elif hours > 40:
        wage2 = (40 * wage) + ((hours - 40) * ((1.5) * wage)) 
        return wage2
    

```

<!-- grade_id: cell-635b5067a1408dac, points: 2, grade: true, solution: false -->
```python
# Test cell for exercise 3
assert total_wages_for_the_week(40, 25) == 1000
assert total_wages_for_the_week(65, 63.4) == 4913.5
```

<!-- grade_id: cell-16a48177c3391f17, points: 2, grade: true, solution: false -->
```python
# Test cell for exercise 3 (Hidden)
```

<!-- grade_id: cell-4537309f8b40b433, points: 0, grade: false, solution: true -->
```python
# Exercise 4
def convert_quiz_score_to_letter_grade(score):
    if score == 0 or score == 1:
        return "F"
    #elif score == 1:
     #   print("F")
    elif score == 2:
        return "D"
    elif score == 3:
        return "C"
    elif score == 4:
        return "B"
    elif score == 5:
        return "A"
    else:
        return "Not a valid score"
    
```

<!-- grade_id: cell-59258c70892bdfe2, points: 4, grade: true, solution: false -->
```python
# Test cell  for exercise 4 
assert convert_quiz_score_to_letter_grade(0) == "F"
assert convert_quiz_score_to_letter_grade(1) == "F"
assert convert_quiz_score_to_letter_grade(2) == "D"
assert convert_quiz_score_to_letter_grade(3) == "C"
assert convert_quiz_score_to_letter_grade(4) == "B"
assert convert_quiz_score_to_letter_grade(5) == "A"
```

<!-- grade_id: cell-938cf03b8db83592, points: 0, grade: false, solution: true -->
```python
# Exercise 5 
def convert_exam_score_to_letter_grade(score):
    # A certain CS professor gives 100-point exams that are graded on the scale 90-100: A, 80-89: B, 70-79: C,
    # 69-69: D, < 60: F. Write a function that accepts an exam score as input and uses a decision structure 
    # to calculate the corresponding grade. 
    # YOUR CODE HERE
    if score >= 90 and score <= 100:
        return "A"
    elif score >= 80 and score <= 89:
        return "B"
    elif score >= 70 and score <= 79:
        return "C"
    elif score >= 60 and score <= 69:
        return "D"
    elif score < 60:
        return "F"
    
```

<!-- grade_id: cell-913bb7905082e4f6, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 5
assert convert_exam_score_to_letter_grade(55) == "F"
assert convert_exam_score_to_letter_grade(35) == "F"
assert convert_exam_score_to_letter_grade(65) == "D"
assert convert_exam_score_to_letter_grade(73) == "C"
assert convert_exam_score_to_letter_grade(87) == "B"
assert convert_exam_score_to_letter_grade(92) == "A"
```

<!-- grade_id: cell-c0e4ea4c0dd98cb5, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 5 (Hidden)
```

<!-- grade_id: cell-87e5ab1727332fc9, points: 0, grade: false, solution: true -->
```python
# Exercise 6
def calculate_class_standing_from_credits(credits):
    # A certain college classifies students according to credits earned. A student with less than 7 credits is a Freshman.
    # At least 7 credits are required to be a Sophomore, 16 to be a Junior and 26 to be classifed as Senior. 
    # Write a function that calculates class standing from teh number of credits earned. 
    # YOUR CODE HERE
    if credits < 7:
        return "Freshman" 
        #print("Freshman")
    elif credits >= 7 and credits < 16:
        return "Sophomore" 
        #print("Sophomore")
    elif credits >= 16 and credits < 26:
        return "Junior" 
        #print("Junior")
    elif credits >= 26:
        return "Senior" 
        #print("Senior")
    
```

<!-- grade_id: cell-b8f9a22bb432f494, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 6
assert calculate_class_standing_from_credits(6) == "Freshman"
assert calculate_class_standing_from_credits(13) == "Sophomore"
assert calculate_class_standing_from_credits(23) == "Junior"
assert calculate_class_standing_from_credits(45) == "Senior"

```

<!-- grade_id: cell-4f3ca68adff20e03, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 6
```

<!-- grade_id: cell-69316b30d8c33c42, points: 0, grade: false, solution: true -->
```python
# Exercise 7
def calculate_bmi(weight, height):
    # The body mass index (BMI) is calculated as a peron's weight (in pounds) times 720, divided by the square
    # of the person's height (in inches). A BMI in teh rnage 19-25, inclusive, is considered health. Write a function
    # that calculates a person's BMI and prints a message telling whether they are above, within, or below the 
    # the healthy range. 
    # Round to two decimals places
    heightsq = height**2
    bmi = (weight * 720) / heightsq
    if bmi >= 19 and bmi <= 25:
        print("Healthy")
        return round(bmi,2)
    elif bmi > 25:
        print("Above healthy range")
        return round(bmi,2)
    elif bmi < 19:
        print("Below healthy range")
        return round(bmi,2)
    # YOUR CODE HERE
    raise NotImplementedError()
    
    
```

<!-- grade_id: cell-b2edb2846eaa98a5, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 7
assert calculate_bmi(150, 65) == 25.56
```

Output:
```
Above healthy range
```

<!-- grade_id: cell-ad3c15e61f39cd2a, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 7 (hidden)
```

<!-- grade_id: cell-458b93f22ceefd76, points: 0, grade: false, solution: true -->
```python
# Exercise 8
def calculate_fine(limit, speed):
    # The speeding ticket fine policy in Podunksville is $50 plus $5 for each mph over the limit plus a 
    # penalty of $200 for any speed over 90 mph. Write a function that accepts a speed limit and a clocked
    # speed and either return the string `Legal` or the amount of fine, if the speed is illegal. 
    # The speeding ticket fine policy in Podunksville is $50 plus $5 for each mph over the limit plus a 
    # penalty of $200 for any speed over 90 mph. Write a function that accepts a speed limit and a clocked
    # speed and either return the string `Legal` or the amount of fine, if the speed is illegal. 
    if speed > limit and speed <= 90:
        fine = 50 + ((speed - limit)*5)
        return fine
    if speed > 90:
        fine1 = 250 + ((speed - limit)*5) 
        return fine1
    elif speed <= limit:
        return "Legal"
    
```

<!-- grade_id: cell-0e86617dd30776cc, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 8
assert calculate_fine(55, 85) == 200
assert calculate_fine(55, 45) == 'Legal'
assert calculate_fine(55, 100) == 475
```

<!-- grade_id: cell-3aa05ca45ba9b9ba, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 8
```

<!-- grade_id: cell-9a1f652136c7a757, points: 0, grade: false, solution: true -->
```python
# Exercise 9
def is_leap_year(year):
    # A year is a leap year if it is divisible by 4, unless it is a century year that is not divisible 
    # by 400. (1800 and 1900 are not leap years while 1600 and 2000 are). Write a function that calculates
    # whether a year is a leap year. Return True or False

    if year % 4 == 0:
        if year % 100 == 0:
            if year % 400 == 0:
                #print("True")
                return True
            else:
                #print("False")
                return False
        else:
             return True
    else:
        return False

```

<!-- grade_id: cell-a64558eb5104a70a, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 9
assert is_leap_year(2000) == True
assert is_leap_year(1800) == False
assert is_leap_year(1992) == True
```

<!-- grade_id: cell-e720e02f8301a7c3, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 9
```

<!-- grade_id: cell-4f0751ced1b247b5, points: 0, grade: false, solution: true -->
```python
# Exercise 10
def is_palindrome(input_string):
    # Write a function that determines if the input_string is a palindrome. Return True or False
    # Use square brackets to reverse the input_string! Make sure to lower the input string before testing!
    inputstring = input_string.lower()
    if (inputstring == inputstring[::-1]):
        return True
    else:
        return False

```

<!-- grade_id: cell-b317cf4721c1c02b, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 10
assert is_palindrome("Kayak") == True
assert is_palindrome("Rotator") == True
assert is_palindrome("AACA") == False

```

<!-- grade_id: cell-fb9df3154f5e8cad, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 10 (hidden)
```

```python

```