Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Abinesh Senthil Kumar"
COLLABORATORS = ""
```

---

<!-- grade_id: cell-a4ddc5d59b1c64bb, points: 0, grade: false, solution: true -->
```python
# Ex1
def average_word_len_in_sentence(sentence):

    splitted = sentence.split(" ")
    length = len(splitted)
    eachlen = 0
    for x in splitted:
        eachlen += len(x)
    
    return round((eachlen/length),2)
```

<!-- grade_id: cell-014af2cf28c5fac7, points: 2, grade: true, solution: false -->
```python
assert average_word_len_in_sentence('This is a test sentence!') == 4
```

<!-- grade_id: cell-251f809b4f7d4e57, points: 2, grade: true, solution: false -->
```python
# Hidden test

```

<!-- grade_id: cell-4afbc828a658cb2a, points: 0, grade: false, solution: true -->
```python
#Ex2    
def wc(filename):

    file = open(filename, 'r')

    number_lines = 0
    number_words = 0
    number_char = 0

    for x in file:
        number_lines += 1 
        length_char = x.split()
        number_words += len(length_char)
        number_char += len(x)
    
    return (number_lines, number_words, number_char)
    

```

<!-- grade_id: cell-9449925bf7f9a740, points: 4, grade: true, solution: false -->
```python
assert wc('ex2_data.txt') == (50, 643, 4298)

```

<!-- grade_id: cell-44c50229f0f432b1, points: 0, grade: false, solution: true -->
```python
# Ex3a
def is_leap(year):
    if year % 4 == 0:
        if year % 100 == 0:
            if year % 400 == 0:
                #print("True")
                return True
            else:
                #print("False")
                return False
        else:
             return True
    else:
        return False

```

<!-- grade_id: cell-3f56767fb1829b44, points: 1, grade: true, solution: false -->
```python
assert is_leap(1800) == False
assert is_leap(1900) == False
assert is_leap(1600) == True
assert is_leap(2000) == True
```

<!-- grade_id: cell-ec2a09f5b6a54636, points: 0, grade: false, solution: true -->
```python
#Ex3b
def is_date_valid(month, day, year):
   
    def is_leap_year(year):
    
        if year % 4 == 0:
            if year % 100 == 0:
                if year % 400 == 0:
                
                    return True
                else:
                
                    return False
            else:
                  return True
        else:
            return False
    
    day = int(day)
    month = int(month)
    year = int(year)
    
    month_date_dict = {1:range(1,32), 2:range(1,29), 3:range(1,32), 4:range(1,31),
                       5:range(1,32), 6:range(1,31), 7:range(1,32), 8:range(1,32),
                       9:range(1,31), 10:range(1,32), 11:range(1,31), 12:range(1,32)}
    
    if is_leap_year(year) == True:
        if month in range(1,13):
            if month == 2:
                if day in range(1,30):
                    return True
                else:
                    return False
            elif day in month_date_dict[month]:
                return True
            else:
                return False
        else:
            return False
       
    if is_leap_year(year) == False:
        if month in range(1,13):
            if day in month_date_dict[month]:
                return True
            else:
                return False
        else:
            return False
    
```

<!-- grade_id: cell-2351e73237cb403c, points: 1, grade: true, solution: false -->
```python
assert is_date_valid(5, 24, 1962) == True
assert is_date_valid(9, 31, 2000) == False
```

<!-- grade_id: cell-fe0d25010ddcf60b, points: 2, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-deffc09fbb6e4ff0, points: 0, grade: false, solution: true -->
```python
#Ex4            
def day_number(month, day, year):   
    
    def is_date_valid(month, day, year):
   
        def is_leap_year(year):
    
            if year % 4 == 0:
                if year % 100 == 0:
                    if year % 400 == 0:               
                        return True
                    else:               
                        return False
                else:
                    return True
            else:
                return False
    
        day = int(day)
        month = int(month)
        year = int(year)
    
        month_date_dict = {1:range(1,32), 2:range(1,29), 3:range(1,32), 4:range(1,31),
                           5:range(1,32), 6:range(1,31), 7:range(1,32), 8:range(1,32),
                           9:range(1,31), 10:range(1,32), 11:range(1,31), 12:range(1,32)}
    
        if is_leap_year(year) == True:
            if month in range(1,13):
                if month == 2:
                    if day in range(1,30):
                        return True
                    else:
                        return False
                elif day in month_date_dict[month]:
                    return True
                else:
                    return False
            else:
                return False
       
        if is_leap_year(year) == False:
            if month in range(1,13):
                if day in month_date_dict[month]:
                    return True
                else:
                    return False
            else:
                return False
            
    if is_date_valid(month, day, year) == True:
    
        import datetime as dt

        d0 = dt.date(year,1,1)
        d1 = dt.date(year, month, day)


        diff = (d1 - d0)

        return (diff.days+1)
        
    else:
        
        return False


```

<!-- grade_id: cell-d6eaeeca3283f71c, points: 2, grade: true, solution: false -->
```python
assert day_number(9, 31, 2000) == False
assert day_number(2, 13, 2020) == 44
```

<!-- grade_id: cell-4e3c036b2a757130, points: 2, grade: true, solution: false -->
```python
# Hidden tests

```

<!-- grade_id: cell-8ed9233092ac6714, points: 0, grade: false, solution: true -->
```python
#Ex5
def years_to_double_investment(apr):
    principle = 1
    count = 0
    total = principle
    #value = principle
    while total < 2*principle:
        total = total * (1 + apr)
        count += 1
        #total1.append(total)
        
    return (count)
```

<!-- grade_id: cell-9723aeff061e58f8, points: 4, grade: true, solution: false -->
```python
assert years_to_double_investment(0.06) == 12
assert years_to_double_investment(0.1) == 8
```

<!-- grade_id: cell-21c962647acaa4f7, points: 0, grade: false, solution: true -->
```python
#Ex6 
def stopping_time(n):
    n = int(n)
    count = 0
    while n > 1:
        if n % 2 == 0:
            n = n/2
        else:
            n = (n*3) + 1
        
        count += 1
    return (count)
```

<!-- grade_id: cell-f491a588fd348765, points: 2, grade: true, solution: false -->
```python
assert stopping_time(12) == 9
assert stopping_time(27) == 111
```

<!-- grade_id: cell-262eecd85e894adc, points: 2, grade: true, solution: false -->
```python
# hidden tests
```

<!-- grade_id: cell-aca81ecb43fae85b, points: 0, grade: false, solution: true -->
```python
#Ex7
def is_prime(n):
    

    import math as m
    rangeof = range(2,((m.ceil(m.sqrt(n)))+1))
    if n == 1:
        return False
    elif n == 2:
        return False
    elif n == 3:
        return True
    elif n == 4:
        return False
    elif n == 5:
        return True
    elif n == 6:
        return False
    elif n == 7:
        return True
    elif n == 8:
        return False
    elif n == 9:
        return False
    else:
        for i in rangeof:
            if n % i == 0:
                return False
                break            
        return True

```

<!-- grade_id: cell-0b96e1ea0ae047b3, points: 4, grade: true, solution: false -->
```python
assert is_prime(2) == False
assert is_prime(3) == True
assert is_prime(5) == True
assert is_prime(25) == False

```

<!-- grade_id: cell-dac6f81c3067bb57, points: 0, grade: false, solution: true -->
```python
#Ex8
def all_primes(n):
    listf = []
    
    import math as m
    def is_prime(n):
    

        import math as m
        rangeof = range(2,((m.ceil(m.sqrt(n)))+1))
        if n == 1:
            return False
        elif n == 2:
            return False
        elif n == 3:
            return True
        elif n == 4:
            return False
        elif n == 5:
            return True
        elif n == 6:
            return False
        elif n == 7:
            return True
        elif n == 8:
            return False
        elif n == 9:
            return False
        else:
            for i in rangeof:
                if n % i == 0:
                    return False
                    break            
            return True

              
    for num in range(1,n+1):
        if is_prime(num) == True:
            listf.append(num)
    return listf

```

<!-- grade_id: cell-da4a1c19327e9780, points: 4, grade: true, solution: false -->
```python
assert all_primes(5) == [3, 5]
assert all_primes(25) == [3, 5, 7, 11, 13, 17, 19, 23]

```

<!-- grade_id: cell-73cc6f55ad21f007, points: 0, grade: false, solution: true -->
```python
#Ex9
def gcd(m,n):
    while m!= 0:

        temp = m
        m = n%m
        n = temp
        
    return n


```

<!-- grade_id: cell-9d49980644cc48c9, points: 2, grade: true, solution: false -->
```python
assert gcd(25,75) == 25
assert gcd(3,13) == 1
```

<!-- grade_id: cell-7ee1b35ba4536fe5, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-70099f7d154d3a41, points: 0, grade: false, solution: true -->
```python
#Ex10
def determine_min_max_average(filename):
    file = open(filename, 'r')

    linelist = file.readlines()

    linelist = [x.strip() for x in linelist]
    
    linelist = [line for line in linelist if line.strip()]
    
    linelist = [i.split(",") for i in linelist]
    
    marksonly = [x[1:] for x in linelist]    
    
    studentname = [x[0] for x in linelist]

    avg = []

    marksonly = [[int(y) for y in x]for x in marksonly]

    avg = []
    for x in marksonly:
        avg.append(sum(x)/len(x))
        
    avg.index(max(avg))
    avg.index(min(avg))

    tupled = (studentname[avg.index(max(avg))] , max(avg) , studentname[avg.index(min(avg))] , min(avg))
    
    return tupled
```

<!-- grade_id: cell-09e42883bbf8750a, points: 4, grade: true, solution: false -->
```python
assert determine_min_max_average('ex10_data.txt') == ('student733', 86.2, 'student202', 65.4)
```

```python

```

```python

```

```python

```