Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Abinesh Senthil Kumar"
COLLABORATORS = ""
```

---

# Ex 1 - Handling Exceptions (10 Points)
In this problem you will read from a text file (**data_p2.txt**) included in the assignment. Please download it and store it in the same directory as this notebook. The code for reading data from the file is given below. You need to store the data and perform a simple mathematical operation on each data point. Each line is _supposed_ to contain a floating point number. But what you will observe is that some lines might have erroneous entries. You need to ignore those lines (Hint: Use `Exception` handling). 

The idea is to implement a function, `computeMedian()`, which reads in a file and computes the median of the numbers and returns the output. You may use the inbuilt function `sort` when computing the median.

_DO NOT USE ANY INBUILT OR OTHER FUNCTION TO DIRECTLY COMPUTE MEDIAN_

<!-- grade_id: cell-bedba72326b592e6, points: 0, grade: false, solution: true -->
```python
# Reading from a file. The following code just reads in every line in a text file without doing anything with it.
# You can use the code in your solution. DO NOT CALL this function in your implementation.
def dummyFileReader():
    listu = []
    with open('data_p2.txt','r') as f: #we open the file in 'read' mode. The 'with' clause is similar to 'finally' clause
        for line in f: #iterate over the file line by line
            s = line.strip() #strip() removes the endline character at the end of the line. Line is of type 'str'
            listu.append(s)
    return listu
            
def computeMedian():
    ad = dummyFileReader()
    newly = []
    for i in ad:
        try:
            #float(i)
            newly.append(float(i))
        except:
            continue
            
    #ad = newly
    newly.sort()
    #print(newly)
        
    if len(newly)%2 == 0:
        #print(len(ad)/2)
        #print(float(ad[int(len(ad)/2)]))
        median = (float(newly[int(len(newly)/2)]) + float(newly[int(len(newly)/2)+1]))/2
        print(median)
        return median
    
    else:
        return (newly[int(len(newly)/2)])
    """
    The function has no input parameter, return median value of the list
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-f8dcb1016008c0d8, points: 10, grade: true, solution: false -->
```python
assert computeMedian() == 0.499675
```

# Ex 2 - The Two Envelopes Problem (20 Points)
The Two Envelopes problem is an interesting decision theory problem.

> Consider that you are given two indistinguishable envelopes, each containing money, one contains twice as much as the other. You may pick one envelope and keep the money it contains. Having chosen an envelope at will, but before inspecting it, you are given the chance to switch envelopes. Should you switch?

> The game is this: "stick" or "switch"; It seems obvious that there is no point in switching envelopes as the situation is symmetric. However, because you stand to gain twice as much money if you switch while risking only a loss of half of what you currently have, it is possible to argue that it is more beneficial to switch. The problem is to show what is wrong with this argument.

**Two Envelopes problem**: Implement a function, called `simulateProblem()`, that does the game simulation for the two envelopes problem. Run the simulation 10000 times to figure out the empirical (observed) probability of gaining more money when switching and gaining more money when sticking to the original choice. Each simulation operates as follows:
1. First, randomly pick an envelopes configuration out of the two possible configurations, $(A,2A)$ or $(2A,A)$. In the first configuration, the second envelope has twice the money and in the second configuration, the first envelope has twice the money.
2. Next, randomly pick one of the two envelopes. 
3. Finally, randomly choose to either stick or switch. The program checks if you won (the envelope that picked has more money) or not (the envelope that picked has less money). In case of winning, record if the winning was because of _sticking_ or _switching_.

You can perform the _random_ choice as follows, using the `np.random.randint()` method.
```python
import numpy as np
print(np.random.randint(2))
```
The `simulateProblem()` function takes no arguments and returns two values, first is a boolean output which is `True` if you win and `False` if you lose. In case of a win, the second output is `True` if the win was due to _sticking_ or the lose was due to _switching_ and `False` if the win was due to _switching_ or the lose was due to _sticking_.

Once the method `simulateProblem()` that does the above steps and returns "sticking",or "switching", depending on the win/loss scenario, run the method 1000 times and count the number of times the win was due to _sticking_ to the pick in Step 2, and number of times the win was due to _switching_ the envelope.

<!-- grade_id: cell-9d6f62b5fca5a02a, points: 0, grade: false, solution: true -->
```python
import numpy as np

def simulateProblem():
    # YOUR CODE HERE
    a = (np.random.randint(2))
    #b = (np.random.randint(2))
    
    if a == 1:
        first = 100
        second = 50
        #print(first,second)
    elif a == 0:
        first = 50
        second = 100
        #print(first,second)
    
    pick = np.random.randint(2)
    if pick == 0:
        userpicked = first
        #print("picked first")
    else:
        userpicked = second
        #print("picked second")
    stickorswitch = np.random.randint(2)
    if stickorswitch == 0:
        a = "sticking"
        #print("sticked with what i had")
        pass
    else:
        if userpicked == first:
            userpicked = second
            a = "switching"
            #print("switched to second")
        else:
            userpicked == first
            a = "switching"
            #print("switched to first")
            
        
    if userpicked > first:
        #print("win")
        if a == "sticking":
            return (True, True)
        if a == "switching":
            return (True, False)
        
    elif userpicked > second:
        #print("win")
        if a == "sticking":
            return (True, True)
        if a == "switching":
            return (True, False)
        
    else:
        #print("lost")
        if a == "switching":
            return (False,True)
        if a == "sticking":
            return (False,False)
        
        
    
    
    #raise NotImplementedError()
```

<!-- grade_id: cell-90a1674a5478ea01, points: 5, grade: true, solution: false -->
```python
assert type(simulateProblem()[0]) == bool
assert type(simulateProblem()[1]) == bool
```

<!-- grade_id: cell-53333d2767feff38, points: 0, grade: false, solution: true -->
```python
def run_simulation():
    """
    The function Run the simulation 10000 times to figure out 
    the empirical (observed) probability of gaining more money when switching 
    and gaining more money when sticking to the original choice.
    Return the probability of win due to sticking and win due to switching
    """
    
    winduetostick = 0
    winduetoswitch = 0
    
    for i in range(10000):
        ans = simulateProblem()
        if ans[0] == True:
            if ans[1] == True:
                winduetostick += 1
            elif ans[1] == False:
                winduetoswitch += 1
                
#     print(winduetoswitch/10000)
#     print(winduetostick/10000)
    return (winduetostick/10000,winduetoswitch/10000)
    # YOUR CODE HERE
#     raise NotImplementedError()
```

<!-- grade_id: cell-58e223a147cfcd13, points: 10, grade: true, solution: false -->
```python

win_due_to_sticking, win_due_to_switching = run_simulation()
assert(abs(win_due_to_sticking-0.25) < 0.05)
assert(abs(win_due_to_switching-0.25) < 0.05)
```

<!-- grade_id: cell-25b2c8ef33953045, points: 0, grade: false, solution: true -->
```python
# EX3

class IteratorClass:
    def __init__ (self,list1,list2,op):
        if type(list1) == list or type(list1) == tuple:
            self.list1 = list1
        else:
            raise ValueError
            
        if type(list2) == list or type(list2) == tuple:
            self.list2 = list2
        else:
            raise ValueError
        if str(op) == "add" or str(op) == "sub" or str(op) == "mul" or str(op) == "div":
            self.op = op
        else:
            raise ValueError 
            
    def __iter__(self):
        self.index = 0    
        
        return self
    
    def __next__(self):
        if self.index < len(self.list1) and self.index < len(self.list2):
            
            if self.op =="add":
                output = self.list1[self.index] + self.list2[self.index]
#                 print("index:",self.index)
                self.index += 1
            if self.op =="sub":
                output = self.list1[self.index] - self.list2[self.index]
                self.index += 1
            if self.op =="mul":
                output = self.list1[self.index] * self.list2[self.index]
                self.index += 1
            if self.op =="div":
                output = round((self.list1[self.index] / self.list2[self.index]),2)
                self.index += 1
                    
            return output
        else:
            raise StopIteration
    
    
    

    
    
    
    
    
    
    
    
    
    
    
    
    
    
    
    
#     def __init__(self, list1, list2, operator):
        
#         self.list1 = list1
#         self.list2 = list2
#         self.operator = operator
#         self.limit = len(list1)
        
    
#     def listchecker(self):
        
#         if type(self.list1) is list or type(self.list1) is tuple:
#             pass
#         else:
#             raise ValueError("input1 is not list or tuple")
            
#         if type(self.list2) is list or type(self.list2) is tuple:
#             pass
#         else:
#             raise ValueError("input2 is not list or tuple")
#         # try:
#         #     list(self.list1)
#         #     try:
#         #         tuple(self.list1)
#         #     except ValueError:
#         #         raise ValueError("input1 is not list")
#         # except:
#         #     raise ValueError("input1 is not list")
        
#         # try:
#         #     list(self.list2)
#         #     try:
#         #         tuple(self.list2)
#         #     except ValueError:
#         #         raise ValueError("input2 is not list")
#         # except:
#         #     raise ValueError("input2 is not list")
            
#     def operatorchecker(self):
        
#         if str(self.operator) == "add" or str(self.operator) == "sub" or str(self.operator) == "mul" or str(self.operator) == "div":
#             pass
#         else:
#             raise ValueError("enter valid operator")
    
#     def listoperations(self):
        
#         if self.operator == "add":
#             self.added = [x+y for x,y in zip(self.list1,self.list2)]
            
#         if self.operator == "sub":
#               self.subtract = [x-y for x,y in zip(self.list1,self.list2)]
              
#         if self.operator == "mul":
#             self.multiplication = [x*y for x,y in zip(self.list1,self.list2)]
            
#         if self.operator == "div":
#             self.division = [x/y for x,y in zip(self.list1,self.list2)]
    
#     def __iter__(self):
#         self.value = 0
#         return self
    
#     def __next__(self):
        
        
#         if self.value < self.limit:
        
            
#             self.listchecker()
#             self.operatorchecker()
#             self.listoperations()
        
#             if self.operator == "add":
#                 self.value += len(self.list1)
#                 return (self.added)
                
        
#             if self.operator == "sub":
#                 self.value += len(self.list1)
#                 return (self.subtract)
                
        
#             if self.operator == "mul":
#                 self.value += len(self.list1)
#                 return (self.multiplication)
        
#             if self.operator == "div":
#                 self.value += len(self.list1)
#                 return (self.division)
                
                
#         else:
#             raise StopIteration 
    
#     ## Complete this class! It takes in three inputs when initializing. 
#     # input#1 x -- is a sequence, either a list or a tuple. Raise a ValueError if it is not a list or a tuple
#     # input#2 y -- is a sequence, either a list or a tuple. Raise a ValueError if it is not a list or a tuple
#     # input#3 operator -- is a string that can either be 'add', 'sub', 'mul', 'div' -- If the specified operator
#     # is not one of these, raise a ValueError. 
    
#     # Complete the class by writing functions that will turn it into an iterator class. 
#     # https://www.programiz.com/python-programming/methods/built-in/iter
#     # The purpose of the class is to take two lists(x and y), apply the specified operator and return the output
#     # as an iterator, meaning you can do "for ele in IteratorClass(x,y,'add')"
    
    
#     # YOUR CODE HERE
#     #raise NotImplementedError()
    
    

```

<!-- grade_id: cell-e6fd6bbad085396e, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_raises
assert_raises(ValueError, IteratorClass, range(5,10), [1,2,3], 'div')
assert_raises(ValueError, IteratorClass, [1,2,3], range(5,10), 'div')
assert_raises(ValueError, IteratorClass, [1,2,3], [1,2,3], 'abc')
```

<!-- grade_id: cell-b148b6133277ed24, points: 10, grade: true, solution: false -->
```python
x = list(range(5,10))
y = list(range(50,55))

add_iterator = IteratorClass(x,y,'add')
sub_iterator = IteratorClass(x,y,'sub')
mul_iterator = IteratorClass(x,y,'mul')
div_iterator = IteratorClass(x,y,'div')

assert [ele for ele in add_iterator] == [55, 57, 59, 61, 63]
assert [ele for ele in sub_iterator] == [-45, -45, -45, -45, -45]
assert [ele for ele in mul_iterator] == [250, 306, 364, 424, 486]
assert [ele for ele in div_iterator] == [0.1, 0.12, 0.13, 0.15, 0.17]
```

```python

```