Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Abinesh Senthil Kumar"
COLLABORATORS = ""
```

---

<!-- grade_id: cell-e7288cd6b4f835dc, points: 0, grade: false, solution: true -->
```python
# Ex1: Season from Month and Day
def determine_season(month, day):
    month = str(month.lower())
    day = int(day)
    
    #months = {1:january,2:february,3:march,4:april,5:may,6:june,7:july,8:august,9:september,10:october,11:november,12:december}
    
    months = {"january":1,"february":2, "march":3, "april":4, "may":5,"june":6,"july":7,"august":8,"september":9,"october":10,"november":11,"december":12}

    monthnum = months[month]

    import datetime as dt
    d0 = dt.date(2019,1,1)
    d1 = dt.date(2019,monthnum,day)
    
    diff = d1 - d0
    
    diff = diff.days
    #print(diff)    
    #d1 = dt.date(2019,12,27)

    if diff >= 78 and diff < 171:
        return "Spring"
    
    if diff >= 171 and diff < 264:
        return "Summer"
    
    if diff >= 264 and diff < 354:
        return "Fall"
    
    if diff >= 354 or diff < 78:
        return "Winter"


```

<!-- grade_id: cell-41b01563ca7bfa78, points: 2, grade: true, solution: false -->
```python
assert determine_season('March', 21) == 'Spring'
assert determine_season('June', 21) == 'Summer'
assert determine_season('November', 21) == 'Fall'
assert determine_season('January', 21) == 'Winter'
```

<!-- grade_id: cell-eee77ee3863f3c90, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-91cc8d6f0acf4f9f, points: 0, grade: false, solution: true -->
```python
# Ex2:Is a License Plate Valid?

def is_license_plate_valid(plate):
    num = str(plate)
    
    num = [i.split() for i in num]
    
    if len(num) == 6:
        
        try:
            num[3][0] = float(num[3][0])
            num[4][0] = float(num[4][0])
            num[5][0] = float(num[5][0])
        except:
            return "Invalid"
        
        if num[0][0] == num[0][0].upper() and num[1][0] == num[1][0].upper() and num[2][0] == num[2][0].upper() and num[3][0].is_integer() and num[4][0].is_integer() and num[5][0].is_integer():
            return "Older/Valid"
        
        else:
            return "Invalid"
    elif len(num) == 7:
        
        try:
            num[0][0] = float(num[0][0])
            num[1][0] = float(num[1][0])
            num[2][0] = float(num[2][0])
            num[3][0] = float(num[3][0])
        except:
            return "Invalid"
        
        if num[0][0].is_integer() and num[1][0].is_integer() and num[2][0].is_integer() and num[3][0].is_integer() and num[4][0] == num[4][0].upper() and num[5][0] == num[5][0].upper() and num[6][0] == num[6][0].upper():
            return "Newer/Valid"
        
        else:
            return "Invalid"
    else:
        return "Invalid"
        

```

<!-- grade_id: cell-a4a6978643ddcfe0, points: 2, grade: true, solution: false -->
```python
assert is_license_plate_valid('ABC123') == 'Older/Valid'
assert is_license_plate_valid('GHE952') == 'Older/Valid'
assert is_license_plate_valid('1934ABT') == 'Newer/Valid'
assert is_license_plate_valid('bTR342') == 'Invalid'
```

<!-- grade_id: cell-46e008a0ccb94026, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-27bf09ccd6919f1a, points: 0, grade: false, solution: true -->
```python
#Ex3: Check a password

def check_password(password):
    
    pwd = password
    
    if len(pwd) < 8:
        return False
    else:
        if any(i.isupper() for i in pwd) == False:
            return False
        if any(i.isdigit() for i in pwd) == False:
            return False
        if any(i.islower() for i in pwd) == False:
            return False
        else:
            return True


```

<!-- grade_id: cell-b3b1e2b92b28aba3, points: 2, grade: true, solution: false -->
```python
assert check_password('test1234') == False
assert check_password('password123') == False
assert check_password('SuperPasswrd90') == True
assert check_password('letmein!') == False

```

<!-- grade_id: cell-9c53951c78f3bcbd, points: 2, grade: true, solution: false -->
```python


```

<!-- grade_id: cell-775a23c9b2859090, points: 0, grade: false, solution: true -->
```python
# Ex4: Magic dates

def is_magic_date(date):
    date = date.split("/")
    
    date = [i.split() for i in date]
    
    date = [int(i[0]) for i in date]
    
    day = date[0]
    
    month = date[1]
    
    year = (str(date[2]))
    
    year = int(year[2:])
    
    if (day*month) == year:
        
        return True
    
    else:
        
        return False
    

```

<!-- grade_id: cell-369c11b23998c1d5, points: 4, grade: true, solution: false -->
```python
assert is_magic_date('6/10/1960') == True
assert is_magic_date('6/8/1948') == True
```

<!-- grade_id: cell-3c0531b8de4d714d, points: 0, grade: false, solution: true -->
```python
def remove_outliers(data, num_outliers):
    
    k = num_outliers
    
    n = range(1,k+1)
    
    for i in n:
        data.pop(data.index(max(data)))

    for i in n:
        data.pop(data.index(min(data)))        

    #len(data)
    result = sorted(data)
    
    return result

```

<!-- grade_id: cell-227155127810eff3, points: 2, grade: true, solution: false -->
```python
import random
random.seed(1234)
data = [random.randint(50, 150) for ele in range(100)]
data[45] = 1
data[46] = 2
data[90] = 250
data[34] = 300

result = [50, 50, 51, 52, 52, 52, 53, 53, 54, 
55, 55, 55, 55, 56, 58, 58, 59, 59, 59, 60, 61, 61, 
61, 62, 64, 64, 64, 68, 68, 68, 69, 69, 69, 70, 71, 
72, 73, 75, 77, 80, 81, 81, 84, 84, 84, 85, 88, 88, 89, 
92, 94, 94, 95, 95, 98, 103, 106, 108, 108, 109, 109, 109, 
110, 111, 111, 112, 113, 114, 115, 115, 117, 117, 119, 121, 
124, 124, 125, 126, 128, 129, 132, 132, 133, 134, 135, 135, 
135, 136, 138, 140, 141, 148, 148, 149, 149, 150]

assert remove_outliers(data, 2) == result

```

<!-- grade_id: cell-4765f8855ee87dc5, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-815256da835e09ed, points: 0, grade: false, solution: true -->
```python
# Ex6: Removing duplicates

def remove_duplicates(words):
    list = []
    for i in words:
        if i not in list:
            list.append(i)
            
    return list
    
    
    
    
    #     for i in words:
#         if words.count(i) >= 2:
#          words.remove(i)   
#     return words

#def remove_duplicates(words):
 #   words = set(words)
  #  return words
    
```

<!-- grade_id: cell-96abefbf1a6700cd, points: 2, grade: true, solution: false -->
```python
assert remove_duplicates([1,2,3,3,3,4,5,6,7,7,8]) == [1, 2, 3, 4, 5, 6, 7, 8]

```

<!-- grade_id: cell-9af188dcb9ade23a, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-12ea8be11e542c51, points: 0, grade: false, solution: true -->
```python
# Ex7: List of proper divisors

def proper_divisors(n):
    range_of_pos = range(1,n)
    prop_div = []
    for i in range_of_pos:
        if n % i == 0:
            prop_div.append(i)
            
    return prop_div
    
    
```

<!-- grade_id: cell-fddcfe9579472263, points: 2, grade: true, solution: false -->
```python
assert proper_divisors(28) == [1, 2, 4, 7, 14]


```

<!-- grade_id: cell-c5d5761e00afab4e, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-08e0b5fe18e178bd, points: 0, grade: false, solution: true -->
```python
#Ex8: Perfect Numbers
def is_number_perfect(n):
    def proper_divisors(n):
        range_of_pos = range(1,n)
        prop_div = []
        for i in range_of_pos:
            if n % i == 0:
                prop_div.append(i)
            
        return prop_div
    
    if sum(proper_divisors(n)) == n:
        return True
    else:
        return False
```

<!-- grade_id: cell-5657e6c68422db39, points: 2, grade: true, solution: false -->
```python
assert is_number_perfect(28) == True
assert is_number_perfect(76) == False

```

<!-- grade_id: cell-a90a7b21c7687d99, points: 2, grade: true, solution: false -->
```python


```

<!-- grade_id: cell-a1ccc2f098ce408c, points: 0, grade: false, solution: true -->
```python
# Ex9: Line of Best Fit
def best_line(points):
    x = []
    y = []
    for i in points:
        
        x.append(i[0])
        y.append(i[1])
    
    avg_x = sum(x)/len(x)
    avg_y = sum(y)/len(y)

    x_minus_xbar = []
    for i in x:
        x_minus_xbar.append(i-avg_x) 

    y_minus_ybar = []
    for i in y:
        y_minus_ybar.append(i-avg_y)    

    product_of_norm_xy = []
    for i,j in zip(x_minus_xbar,y_minus_ybar):
        product_of_norm_xy.append(i*j)    

    x_minus_x_avg_sq = []

    for i in x:
        x_minus_x_avg_sq.append(i-avg_x)

    x_minus_x_avg_sq = [i**2 for i in x_minus_x_avg_sq]


    slope = sum(product_of_norm_xy)/sum(x_minus_x_avg_sq)

    y_intercept = avg_y - (slope*avg_x)

    return (round(slope,2),round(y_intercept,2))


```

<!-- grade_id: cell-f0ac1269eeca97e3, points: 2, grade: true, solution: false -->
```python
points = (
    (8, 3),
    (2, 10),
    (11, 3),
    (6, 6),
    (5, 8),
    (4, 12),
    (12, 1),
    (9, 4),
    (6, 9),
    (1, 14),
)

assert best_line(points) == (-1.11, 14.08)
```

<!-- grade_id: cell-176082640140d417, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-42f4bf8c7b879931, points: 0, grade: false, solution: true -->
```python
# Ex10: Reverse Lookup

def reverse_lookup(input_dict, test_value):

    index = []
    for i in input_dict:
        index.append(i)
        
    values = []
    for i in index:
        values.append(input_dict[i])


    index_no_of_val = []
    for i in range(len(values)):
        if values[i] == test_value:
            index_no_of_val.append(i)

    req_index = []
    for i in index_no_of_val:
        req_index.append(index[i]) 

    return req_index


```

<!-- grade_id: cell-668a78aad343067b, points: 2, grade: true, solution: false -->
```python
input_dict = {
    'January': 31,
    'February': 28,
    'March': 31,
    'April': 30,
    'May': 31,
    'June': 30,
    'July': 31,
    'August': 31,
    'September': 30,
    'October': 31,
    'November': 30,
    'December': 31,
}



assert reverse_lookup(input_dict, 31) == ['January', 'March', 'May', 'July', 'August', 'October', 'December']


```

<!-- grade_id: cell-3f9dbfc14bd930a6, points: 2, grade: true, solution: false -->
```python

```